# Week 4 Problem Set

## Homeworks

**HW1.** Implement `Queue` abstract data structure using a Class. This `Queue` only stores items of Integer data type. You can use a `list` as its internal data structure. The class should have the following interface:
- `__init__()` to initialize an empty List for the queue to store the items.
- `enqueue(item)` which inserts an Integer into the queue.
- `dequeue()` which returns and removes the element at the head of the queue. The return value is an optional as it may return `None` if there are no more elements in the queue.
- `peek()` which returns the element at the head of the queue.

The class Queue has two computed properties:
- `is_empty` which returns either `True` or `False` depending on whether the queue is empty or not.
- `size` which returns the number of items in the queue.

```python
class Queue:
    def __init__(self):
        self.items = []

    def enqueue(self, item):
        self.items.append(item)

    def dequeue(self):
        if self.items == []:
            return None
        else:
            return self.items.pop(0)

    def peek(self):
        if len(self.items) >= 1:
            return self.items[0]

    @property
    def is_empty(self):
        return self.items == []

    @property
    def size(self):
        return len(self.items)
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
q1 = Queue()
q1.enqueue(2)
assert not q1.is_empty
assert q1.size == 1
ans = q1.dequeue()
assert ans == 2
assert q1.is_empty
q1.enqueue(1)
q1.enqueue(2)
q1.enqueue(3)
assert q1.size == 3
assert q1.peek() == 1
assert q1.dequeue() == 1
assert q1.dequeue() == 2
assert q1.dequeue() == 3
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** We are going to create a class that contains both `RobotTurtle` and `Coordinate` class. The name of the class is `TurtleWorld` which is used to simulate when `RobotTurtle` is moving around some two dimensional space. The class has the following methods:

- `add_turtle(name)` which is to add a new `RobotTurtle` into the world with the specified name.
- `remove_turtle(name)` which is to remove the object `RobotTurtle` with the specified name from the world. 
- `list_turtles()` which is to list all the turtles in the world using their names in an ascending order.

We give you here the class definition for the `Coordinate` and the `RobotTurtle` from the Notes.

```python
import math

class Coordinate:
    
    def __init__(self, x=0, y=0):
        self.x = x
        self.y = y
        
    @property
    def distance(self):
        return math.sqrt(self.x * self.x + self.y * self.y)
    
    def __str__(self):
        return f"({self.x}, {self.y})"
```

```python
# Class definition
class RobotTurtle:
    # Attributes:
    def __init__(self, name, speed=1):
        self.name = name
        self.speed = speed
        self._pos = Coordinate(0, 0)
        
    # property getter
    @property
    def name(self):
        return self._name
    
    # property setter
    @name.setter
    def name(self, value):
        if isinstance(value, str) and value != "":
            self._name = value
            
    # property getter
    @property
    def speed(self):
        return self._speed
    
    # property setter
    @speed.setter
    def speed(self, value):
        if isinstance(value, int) and value > 0:
            self._speed = value

    # property getter
    @property
    def pos(self):
        return self._pos
    
    # Methods:
    def move(self, direction):
        update = {'up' : Coordinate(self.pos.x, self.pos.y + self.speed),
                  'down' : Coordinate(self.pos.x, self.pos.y - self.speed),
                  'left' : Coordinate(self.pos.x - self.speed, self.pos.y),
                  'right' : Coordinate(self.pos.x + self.speed, self.pos.y)}
        self._pos = update[direction]

        
    def tell_name(self):
        print(f"My name is {self.name}")

```

Now fill in the class definition for `TurtleWorld`. You may want to look into the test cases in the following cell to make sure you define the class properly.

```python
class TurtleWorld:

    def __init__(self):
        self.turtles = {}

    def add_turtle(self, name, speed):
        self.turtles[name] = RobotTurtle(name, speed)
        # turtle = RobotTurtle(name, speed)
        # self.turtles[name] = turtle

    def remove_turtle(self, name):
        del self.turtles[name]

    def list_turtles(self):
        return sorted([i for i in self.turtles])
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
world = TurtleWorld()
world.add_turtle('t1', 1)
assert world.list_turtles() == ['t1']

world.add_turtle('t2', 2)
assert world.list_turtles() == ['t1', 't2']

world.add_turtle('abc', 3)
assert world.list_turtles() == ['abc', 't1', 't2']

world.remove_turtle('t2')
assert world.list_turtles() == ['abc', 't1']

world.remove_turtle('abc')
assert world.list_turtles() == ['t1']
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** Modify the class `TurtleWorld` to add the following method:
- `move_turtle(name, movement)` which is to move the turtle with the specified name with a given input `movement`. The argument `movement` is a string containing letters: `l` for left, `r` for right, `u` for up, and `d` for down. The movement should be based on the speed. This means that if the turtle has speed of 2 and the `movement` argument is `uulrdd`, the turtle should move up four units, move left two units, move right two units and move down four units.

```python
class TurtleWorld:
    valid_movements = set('udlr')
    movement_map = {'u': 'up', 'd': 'down', 'l': 'left', 'r': 'right'}

    def __init__(self):
        self.turtles = {}

    def move_turtle(self, name, movement): #checkon debug 
        for i in movement:  # i  = TURTLE DIRECTION
            if i in TurtleWorld.valid_movements:  #containing the class under turtleworld 
                self.turtles[name].move(self.movement_map[i])

    def add_turtle(self, name, speed):
        self.turtles[name] = RobotTurtle(name, speed)

    def remove_turtle(self, name):
        del self.turtles[name]

    def list_turtles(self):
        return sorted([i for i in self.turtles])
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
world = TurtleWorld()
world.add_turtle('abc', 1)
world.move_turtle('abc', 'uu')
assert str(world.turtles['abc'].pos) == '(0, 2)'

world.move_turtle('abc', 'rrr')
assert str(world.turtles['abc'].pos) == '(3, 2)'

world.move_turtle('abc', 'd')
assert str(world.turtles['abc'].pos) == '(3, 1)'

world.move_turtle('abc', 'llll')
assert str(world.turtles['abc'].pos) == '(-1, 1)'

world.add_turtle('t1', 2)
world.move_turtle('t1', 'uulrdd')
assert str(world.turtles['t1'].pos) == '(0, 0)'

world.move_turtle('t1', 'ururur')
assert str(world.turtles['t1'].pos) == '(6, 6)'



```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TEST
world = TurtleWorld()
world.add_turtle('dabc', 1)
world.move_turtle('dabc', 'u')
assert str(world.turtles['dabc'].pos) == '(0, 1)'

world.move_turtle('dabc', 'rr')
assert str(world.turtles['dabc'].pos) == '(2, 1)'

world.move_turtle('dabc', 'dd')
assert str(world.turtles['dabc'].pos) == '(2, -1)'

world.move_turtle('dabc', 'llll')
assert str(world.turtles['dabc'].pos) == '(-2, -1)'

world.add_turtle('t1', 2)
world.move_turtle('t1', 'ulrd')
assert str(world.turtles['t1'].pos) == '(0, 0)'

world.move_turtle('t1', 'dldldl')
assert str(world.turtles['t1'].pos) == '(-6, -6)'
###BEGIN HIDDEN TEST
```

**HW4.** Modify the class `TurtleWorld` to include the following method:
- `add_movement(turtle, movement)` which adds turtle movement to a queue to be run later. The argument `turtle` is a string containing the turtle's name. The argument `movement` is another string for the movement. For example, value for `turtle` can be something like `'t1'` while the value for the `movement` can be something like `'uullrrdd'`.
- `run()` which executes all the movements in the queue.

```python
class TurtleWorld:
    valid_movements = set('udlr')
    movement_map = {'u': 'up', 'd': 'down', 'l': 'left', 'r': 'right'}

    def __init__(self):
        self.turtles = {}
        # add a code to create a Queue for the movement
        self._queue_movement = {}

    def add_turtle(self, name, speed):
        self.turtles[name] = RobotTurtle(name, speed)

    def remove_turtle(self, name):
        del self.turtles[name]

    def list_turtles(self):
        return sorted([i for i in self.turtles])

    def add_movement(self, turtle, movement):
        self._queue_movement[turtle] = movement

    def run(self):
        for turtle in self._queue_movement.keys():
            self.move_turtle(turtle, self._queue_movement[turtle])
        self._queue_movement = {}

    def move_turtle(self, name, movement):
        for i in movement:  # i  = TURTLE DIRECTION
            if i in TurtleWorld.valid_movements:  #WHY IS THERE A SELF IN THE VALID MOVEMENTS WHEN IT'S OUTSIDE 
                self.turtles[name].move(self.movement_map[i])

```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
world = TurtleWorld()
world.add_turtle('t1', 1)
world.add_turtle('t2', 2)
world.add_movement('t1', 'ur')
world.add_movement('t2', 'ur')
assert str(world.turtles['t1'].pos) == '(0, 0)'
assert str(world.turtles['t2'].pos) == '(0, 0)'

world.run()
assert str(world.turtles['t1'].pos) == '(1, 1)'
assert str(world.turtles['t2'].pos) == '(2, 2)'

world.run()
assert str(world.turtles['t1'].pos) == '(1, 1)'
assert str(world.turtles['t2'].pos) == '(2, 2)'


```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TEST
world = TurtleWorld()
world.add_turtle('t1', 1)
world.add_turtle('t2', 2)
world.add_movement('t2', 'dldldl')
world.add_movement('t1', 'dldldl')

world.run()
assert str(world.turtles['t1'].pos) == '(-3, -3)'
assert str(world.turtles['t2'].pos) == '(-6, -6)'
###END HIDDEN TEST
```

**HW5.** Implement a radix sorting machine. A radix sort for base 10 integers is a *mechanical* sorting technique that utilizes a collection of bins:
- one main bin 
- 10 digit-bins

Each bin acts like a *queue* and maintains its values in the order that they arrive. The algorithm works as follows:
- it begins by placing each number in the main bin. 
- Then it considers each value digit by digit. The first value is removed from the main bin and placed in a digit-bin corresponding to the digit being considered. For example, if the ones digit is being considered, 534 will be placed into digit-bin 4 and 667 will placed into digit-bin 7. 
- Once all the values are placed into their corresponding digit-bins, the values are collected from bin 0 to bin 9 and placed back in the main bin (in that order). 
- The process continues with the tens digit, the hundreds, and so on. 
- After the last digit is processed, the main bin will contain the values in ascending order.

Create a class `RadixSort` that takes in a List of Integers during object instantiation. The class should have the following properties:
- `items`: is a List of Integers containing the numbers.

It should also have the following methods:
- `sort()`: which returns the sorted numbers from `items` as an `list` of Integers.
- `max_digit()`: which returns the maximum number of digits of all the numbers in `items`. For example, if the numbers are 101, 3, 1041, this method returns 4 as the result since the maximum digit is four from 1041. 
- `convert_to_str(items)`: which returns items as a list of Strings (instead of Integers). This function should pad the higher digits with 0 when converting an Integer to a String. For example if the maximum digit is 4, the following items are converted as follows. From `[101, 3, 1041]` to `["0101", "0003", "1041"]`.

Hint: Your implementation should make use of the generic `Queue` class, which you created, for the bins.

```python
class RadixSort:

    def __init__(self, MyList):
        self.items = MyList

    def max_digit(self):
        max_length = 0
        for i in self.items:
            max_length = max(max_length, len(str(i)))
        return max_length

    def convert_to_str(self, items):
        max_length = self.max_digit()
        for i in range(0, len(items)):
            items[i] = "0" * (max_length - len(str(items[i]))) + str(items[i]) #Put 0's before the digit
        return items

    def sort(self):
        max_length = self.max_digit()
        data_new = self.convert_to_str(self.items)
        container = {}
        for i in range(0, max_length + 1):
            for j in data_new:
                number_to_put_in = int(int(j) % pow(10, i + 1) / pow(10, i))  #Week 4 cohort
                container.setdefault(number_to_put_in, []).append(int(j))
            data_new.clear()
            for i in range(0, 10):
                if container.get(i) != None:
                    for j in container.get(i): #j = value of data
                        data_new.append(int(j))  #Convert back to Integer
            container = {}
        return data_new
```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
list1 = RadixSort([101, 3, 1041])
assert list1.items == [101,3,1041]
assert list1.max_digit() == 4
assert list1.convert_to_str(list1.items) == ["0101", "0003", "1041"]
ans = list1.sort()
print(ans)
assert ans == [3, 101, 1041]
list2 = RadixSort([23, 1038, 8, 423, 10, 39, 3901])
assert list2.sort() == [8, 10, 23, 39, 423, 1038, 3901]
```

Output:
```
[3, 101, 1041]
```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```