# Week 5 Problem Set

## Homework

**HW1.** *Dictionary:* Write two functions:
1. `count_degrees(G)`: which sums up the degrees of all vertices in the graph. The degree of a vertex is defined as the number of edges connected to a Vertex. The argument `G` is a dictionary that represents the graph.
2. `count_edges(G)`: which counts the number of edges in the graph. An edge is defined as a connection between two vertices. The argument `G` is a dictionary.

```python
def count_degrees(G):
    ###
    ### YOUR CODE HERE
    ###
    count = 0
    for key,values in G.items():
        for value in values:
            count += 1
    return count 
    
def count_edges(G):
    ###
    ### YOUR CODE HERE
    ###
    edges = count_degrees(G)/2
    return edges
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
g1 = {"A": ["B", "E"],
      "B": ["A", "C"],
      "C": ["B", "D", "E"],
      "D": ["C"],
      "E": ["A", "C"]}

d = count_degrees(g1)
e = count_edges(g1)

assert d == 10
assert e == 5
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** Create a class called `GraphSearch` which is a subclass of the class `Graph`. This class should override the method `_create_vertex(id)` and instantiate a `VertexSearch` object instead of `Vertex`.

```python
# Copy Vertex class from Cohort
class Vertex:
    ###
    ### YOUR CODE HERE
    ###
    def __init__(self, id=""):
        self.id = id # string
        self.neighbours = {} # dictionary eg: {"Vi_instance": 1, "Vj_instance":3}
    
    # nbr_vertex: vertex instance
    # weight: int 

    def add_neighbour(self, nbr_vertex, weight=0):
        self.neighbours[nbr_vertex] = weight
    
    def get_neighbours(self):
        return list(self.neighbours.keys())
    
    # neighbour: Vertex instance
    def get_weight(self, neighbour):
        return self.neighbours.get(neighbour, None)
    
    def __eq__(self, other):
        return self.id == other.id
    
    def __lt__(self, other):
        return self.id < other.id
    
    def __hash__(self):
        return hash(self.id)
    
    def __str__(self):
        return "Vertex {self.id} is connected to: ".format(self=self) \
                + ", ".join([x.id for x in self.get_neighbours()])

```

```python
# Copy VertexSearch class from Cohort

import sys

class VertexSearch(Vertex):
        ###
        ### YOUR CODE HERE
        ###
    def __init__(self, id=""):
        ###
        ### YOUR CODE HERE
        ###
        super().__init__(id) 
        self.colour = "white"
        self.d = sys.maxsize 
        self._f = sys.maxsize
        self._parent = None 
        
    @property 
    def f(self):
        return self._f

    @f.setter
    def f(self, value):
        self._f = value
    
    @property 
    def parent(self):
        return self._parent
    
    @parent.setter
    def parent(self, value):
        self._parent = value




```

```python
#Copy Graph over from Cohort
class Graph:
    ###
    ### YOUR CODE HERE
    ###
    def __init__(self):
        self.vertices = {} # Dict, key = string of Vertex id, value is Vertex instance
        # to get the neighbouts of a vertex in this graph: self.vertices[vid].neighbours
    def _create_vertex(self, id):
        #create a new vertex object and run it
        return Vertex(id)
    
    def add_vertex(self, id):
        v = self._create_vertex(id)
        # add it to this instance's verties attribute 
        self.vertices[v.id] = v 
    
    def get_vertex(self, id):
        return self.vertices.get(id, None)
    
    def add_edge(self, start_v, end_v, weight=0):
        # check if start_v is in the graph's vertices dict
        if start_v not in self.vertices.keys():
            self.add_vertex(start_v)

        start_vertex = self.vertices[start_v]

        if end_v not in self.vertices.keys():
            self.add_vertex(end_v)

        end_vertex = self.vertices[end_v]

        start_vertex.add_neighbour(end_vertex, weight)
    
    def get_neighbours(self, id):
        v = self.get_vertex(id)
        if v is not None:
            neighbours = v.get_neighbours()
            neighbours_id = []
            for neighbouring_vertex in neighbours:
                # neighbouring_vertex is a vertex instance, so extract its id 
                neighbours_id.append(neighbouring_vertex.id)
            return neighbours_id
        return None
    
    def __contains__(self, val):
        return val in self.vertices.keys()

    def __iter__(self):
        for k,v in self.vertices.items():
            yield v 
        
    # write a code to create a computed property called num_vertices
    @property
    def num_vertices(self):
        return len(self.vertices)
```

```python
class GraphSearch(Graph):
    ##BEGIN SOLUTION
    def _create_vertex(self, id):
        return VertexSearch(id)
    ##END SOLUTION
    pass
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
g2 = GraphSearch()
g2.add_vertex("Z")
assert(type(g2.vertices["Z"]) == type(VertexSearch()))
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *Undirected Graph:* **You need to complete CS5 and CS6 before you do this homework.** Paste your answer for `Search2D` and `SearchBFS` classes into the cell below.

Create a subclass of `GraphSearch` class called `UGraphSearch` for undirected graphs. All edges in undirected graphs are bidirectional (i.e. vertex1 <-> vertex2). 
This means that you need to override the `add_edge()` method.

```python
#Copy Queue over from Wk4 Homework
class Queue:
    pass
    ###
    ### YOUR CODE HERE
    ###

    def __init__(self):
        self.__items = []
    
    def enqueue(self, item):
        self.__items.insert(0, item)
    
    def dequeue(self):
        ###
        ### YOUR CODE HERE
        ###
        remove = self.__items.pop()
        return remove
    
    def peek(self):
        ###
        ### YOUR CODE HERE
        ###
        peek = self.__items[-1]
        return peek
    
    def is_empty(self):
        ###
        ### YOUR CODE HERE
        ###
        length = len(self.__items)
        return length == 0
    
    @property
    def size(self):
        ###
        ### YOUR CODE HERE
        ###
        length = len(self.__items)
        return length 

```

```python
# copy Search2D over from CS5

import sys

class Search2D:
      
    ###
    ### YOUR CODE HERE
    ###
    def __init__(self, g):
        self.graph = g
    
    def clear_vertices(self):
        ###
        ### YOUR CODE HERE
        ###
        for vid, vertex in self.graph.vertices.items():
            vertex.colour = "white"
            vertex.d = sys.maxsize
            vertex.f = sys.maxsize
            vertex.parent = None
    
        pass
    
    def __iter__(self):
        return iter([v for v in self.graph])
    
    def __len__(self):
        return len([v for v in self.graph.vertices])

```

```python
# copy SearchBFS over from CS6

class SearchBFS(Search2D):


    def search_from(self, start):
        ###
        ### YOUR CODE HERE
        ###
        start_vertex = self.graph.get_vertex(start)
        if start_vertex is None:
            return None 
        
        self.clear_vertices()
        
        start_vertex.d = 0
        vertices_to_explore_queue = Queue()
        vertices_to_explore_queue.enqueue(start_vertex)
        
        while not vertices_to_explore_queue.is_empty():
        
            current_vertex_explored = vertices_to_explore_queue.dequeue()
            current_vertex_explored.colour = "black"

        
            for current_vertex_neighbour in current_vertex_explored.get_neighbours():
                if current_vertex_neighbour.colour == "white" and current_vertex_neighbour.parent is None:
                    current_vertex_neighbour.parent = current_vertex_explored
                    current_vertex_neighbour.d = current_vertex_explored.d + 1
                    vertices_to_explore_queue.enqueue(current_vertex_neighbour)
                
    
    def get_shortest_path(self, start, dest):
        pass
        ###
        ### YOUR CODE HERE
        ###
        result = []
        self.get_path(start, dest, result)
        if result == []:
            return None 
        else:
            return result
        
    def get_path(self, start, dest, result):
        ###
        ### YOUR CODE HERE
        ###
        
        start_vertex = self.graph.get_vertex(start)
        if start_vertex is None:
            return None 
        
        end_vertex = self.graph.get_vertex(dest)
        if end_vertex is None:
            return None 

        if start_vertex.d != 0:
            self.search_from(start)

        print(start_vertex.id, end_vertex.id)
        if (start == dest):
            result.append(start)     
        elif end_vertex.parent is None:
            result.append("No Path")
        else:
            self.get_path(start, end_vertex.parent.id, result)
            result.append(dest)
```

```python
class UGraphSearch(GraphSearch):
        ###
        ### YOUR CODE HERE
        ### 
    def add_edge(self, start_v, end_v, weight=0):
        if start_v not in self.vertices:
            self.add_vertex(start_v)
        if end_v not in self.vertices:
            self.add_vertex(end_v)
        
        v_start = self.vertices[start_v]
        v_end = self.vertices[end_v]
        
        v_start.add_neighbour(v_end, weight)
        v_end.add_neighbour(v_start, weight)
        
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
g2 = UGraphSearch()
assert g2.vertices == {} and g2.num_vertices == 0
g2.add_vertex("L")
g2.add_vertex("I")
g2.add_vertex("S")
g2.add_vertex("P")
assert g2.num_vertices == 4
assert "L" in g2
assert "I" in g2
assert "S" in g2
assert "P" in g2
g2.add_edge("L", "I")
g2.add_edge("I", "S")
g2.add_edge("S", "P")
assert sorted(g2.get_neighbours("L")) == ["I"]
assert sorted(g2.get_neighbours("I")) == ["L", "S"]
assert sorted(g2.get_neighbours("S")) == ["I", "P"]
assert sorted(g2.get_neighbours("P")) == ["S"]
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw33, points: 1, grade: true, solution: false -->
```python
g2 = UGraphSearch()
g2.add_vertex("r")
g2.add_vertex("s")
g2.add_vertex("t")
g2.add_vertex("u")
g2.add_vertex("v")
g2.add_vertex("w")
g2.add_vertex("x")
g2.add_vertex("y")
g2.add_vertex("z")
g2.add_edge("r", "s")
g2.add_edge("r", "v")
g2.add_edge("s", "w")
g2.add_edge("t", "u")
g2.add_edge("t", "x")
g2.add_edge("t", "w")
g2.add_edge("u", "t")
g2.add_edge("u", "x")
g2.add_edge("u", "y")
g2.add_edge("v", "r")
g2.add_edge("w", "s")
g2.add_edge("w", "t")
g2.add_edge("w", "x")
g2.add_edge("x", "w")
g2.add_edge("x", "t")
g2.add_edge("x", "u")
g2.add_edge("x", "y")
g2.add_edge("y", "u")
g2.add_edge("y", "x")
gs = SearchBFS(g2)
gs.search_from("s")
assert gs.graph.get_vertex("s").d == 0
assert gs.graph.get_vertex("t").d == 2
assert gs.graph.get_vertex("u").d == 3
assert gs.graph.get_vertex("v").d == 2
assert gs.graph.get_vertex("w").d == 1
assert gs.graph.get_vertex("x").d == 2
assert gs.graph.get_vertex("y").d == 3
assert gs.graph.get_vertex("r").d == 1
ans = gs.get_shortest_path("s", "u")
assert ans == ["s", "w", "t", "u"] or ans == ["s", "w", "x", "u"]
ans = gs.get_shortest_path("v", "s")
assert ans == ["v", "r", "s"]
ans = gs.get_shortest_path("v", "y")
assert ans == ["v", "r", "s", "w", "x", "y"]
```

Output:
```
s u
s t
s w
s s
v s
v r
v v
v y
v x
v w
v s
v r
v v
```

<!-- grade_id: hw34, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *Depth-First-Search:* Create a class `SearchDFS` as a child class of `Search2D` to implement Depth-First-Search algorithm. You should add one additional attribute:
- `time`: which is an attribute to record the discovery time and the finishing time.

The class should also implement the following methods:
- `search()`: which modifies the vertices' properties such as `colour`, `d`, and `parent` following Depth-First-Search algorithm.
- `dfs_visit(vert)`: which is the recursive method for visiting a vertex in Depth-First-Search algorithm.

```python
import sys

class SearchDFS(Search2D):
    def __init__(self, g):
        self.graph = g
        self.time = 0
      
    def search(self):
        ###
        ### YOUR CODE HERE
        ###
        self.clear_vertices()
        
        self.time = 0
        
        for v in self.graph.vertices.values():
            if v.colour == "white":
                self.dfs_visit(v)
    
    def dfs_visit(self, vert):
        ###
        ### YOUR CODE HERE
        ###
        self.time += 1
        vert.d = self.time
        vert.colour = "grey"
    
        for neighbour in self.graph.get_neighbours(vert.id): 
            if self.graph.get_vertex(neighbour).colour == "white":
                self.graph.get_vertex(neighbour).parent = vert
                self.dfs_visit(self.graph.get_vertex(neighbour))
                
        vert.colour = "black"
        self.time +=1
        vert.f = self.time
        
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
g4 = GraphSearch()
g4.add_vertex("e")
g4.add_vertex("m")
g4.add_vertex("a")
g4.add_vertex("c")
g4.add_vertex("s")
g4.add_edge("e", "m")
g4.add_edge("m", "a")
g4.add_edge("a", "c")
g4.add_edge("c", "s")
gs = SearchDFS(g4)
gs.search()
assert gs.graph.get_vertex("e").parent == None 
assert gs.graph.get_vertex("m").parent == gs.graph.get_vertex("e")

assert gs.graph.get_vertex("m").d == 2 and gs.graph.get_vertex("a").f == 8
assert gs.graph.get_vertex("c").d == 4 and gs.graph.get_vertex("s").f == 6
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW5.** *Topological Sort:* Create a function that takes in a `SearchDFS` object to perform a topological sort:
- `topological_sort(g)`: which takes in a `SearchDFS` object and returns a list of `VertexSearch` objects sorted based on their `f` property. This method should call the `search()` method of the `SearchDFS` to first calculate the `f` property of the vertices. Note that you need to copy the `SearchDFS` object of your input argument so as not to mutate the object.

```python
import sys
import copy

def topological_sort(g):
    ls = []
    g.search()

    for value in g.graph.vertices.values():
        ls.append(value)

    ls.sort(key = lambda x : x.f, reverse = True) #checking finishing time 
        
    return ls

    
    
    
```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
import copy
g = GraphSearch()
g.add_vertex("3/4 cup milk")
g.add_vertex("1 egg")
g.add_vertex("1 tbl oil")
g.add_vertex("1 cup mix")
g.add_vertex("heat syrup")
g.add_vertex("heat griddle")
g.add_vertex("pour 1/4 cup")
g.add_vertex("turn when bubbly")
g.add_vertex("eat")
g.add_edge("3/4 cup milk", "1 cup mix")
g.add_edge("1 egg", "1 cup mix")
g.add_edge("1 tbl oil", "1 cup mix")
g.add_edge("1 cup mix", "heat syrup")
g.add_edge("1 cup mix", "pour 1/4 cup")
g.add_edge("heat griddle", "pour 1/4 cup")
g.add_edge("pour 1/4 cup", "turn when bubbly")
g.add_edge("turn when bubbly", "eat")
g.add_edge("heat syrup", "eat")
gs = SearchDFS(g)  

path = topological_sort(gs)
ans = [v.f for v in copy.copy(path)]
assert ans == [18, 16, 14, 12, 11, 10, 9, 6, 5]
ans = [v.id for v in copy.copy(path)]
assert ans == ['heat griddle', '1 tbl oil', '1 egg', '3/4 cup milk', '1 cup mix', 'pour 1/4 cup', 'turn when bubbly', 'heat syrup', 'eat']
```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```